<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Einführung Machine Learning
### Sommersemester 2022
Prof. Dr. Heiner Giefers

<!-- grade_id: cell-4293ef3182378b73, points: 0, grade: false, solution: true -->
```python
measure = [0,0,0,1,0,1,1,0,1,0,0,0,0,0,0,1,0,0,0,0]
predict = [0,0,1,1,0,1,1,0,1,0,0,0,1,0,0,0,0,0,0,0]

TP = None
TN = None
FP = None
FN = None

accuracy  = None
precision = None
recall    = None
f1        = None

# YOUR CODE HERE
raise NotImplementedError()

print("Accuracy = ", accuracy)
print("Precision = ", precision)
print("Recall = ", recall)
print("F1 = ", f1)
```

<!-- grade_id: cell-4cf46ccc45d22173, points: 0, grade: true, solution: false -->
```python
#Test Zelle
### BEGIN HIDDEN TEST
assert abs(accuracy - 17/20) < 1e-6
assert abs(precision - 2/3) < 1e-6
assert abs(recall - 4/5) < 1e-6
assert abs(f1 - 8/11) < 1e-6
### END HIDDEN TEST
```

```python

```